## Autograded Notebook (Canvas & CodeGrade)

This notebook will be automatically graded. It is designed to test your answers and award points for the correct answers. Following the instructions for each Task carefully.

### Instructions

* **Download this notebook** as you would any other ipynb file
* **Upload** to Google Colab or work locally (if you have that set-up)
* **Delete `raise NotImplementedError()`**
* Write your code in the `# YOUR CODE HERE` space
* **Execute** the Test cells that contain `assert` statements - these help you check your work (others contain hidden tests that will be checked when you submit through Canvas)
* **Save** your notebook when you are finished
* **Download** as a `ipynb` file (if working in Colab)
* **Upload** your complete notebook to Canvas (there will be additional instructions in Slack and/or Canvas)

# Unit 1 Sprint 1 Module 2

## EDA and Features

### Objectives

* Load and csv files
* Use basic Pandas function for data exploration
* Understand the purpose of feature engineering

#### Total notebook points: 13

**Task 1** - Load a dataset via its URL

* Create a Python list named `column_headers` with the following items: `'age', 'workclass', 'fnlwgt', 'education', 'education-num', 'marital-status', 'occupation', 'relationship', 'race', 'sex', 'capital-gain', 'capital-loss', 'hours-per-week','native-country', 'income'`.  Lists have the form `listname = ['element 1', 'element 2', element 3']`
* Load the file from the URL provided below using the `column_headers` list you created to name your columns: set the `pd.read_csv()` to `column_headers` using `names =`.  
* Name the DataFrame `adult`

<!-- grade_id: cell-8d0429f74b1e8242, points: 0, grade: false, solution: true -->
```python
# Task 1
import pandas as pd
import numpy as np

# URL for the dataset
data_url = 'https://archive.ics.uci.edu/ml/machine-learning-databases/adult/adult.data'

# YOUR CODE HERE
column_headers = ['age', 'workclass', 'fnlwgt', 'education', 'education-num', 'marital-status', 'occupation', 'relationship', 'race', 'sex', 'capital-gain', 'capital-loss', 'hours-per-week','native-country', 'income']
adult = pd.read_csv(data_url, names = column_headers)

# Print out your DataFrame
print(adult.head())
```

Output:
```
age          workclass  fnlwgt   education  education-num  \
0   39          State-gov   77516   Bachelors             13   
1   50   Self-emp-not-inc   83311   Bachelors             13   
2   38            Private  215646     HS-grad              9   
3   53            Private  234721        11th              7   
4   28            Private  338409   Bachelors             13   

        marital-status          occupation    relationship    race      sex  \
0        Never-married        Adm-clerical   Not-in-family   White     Male   
1   Married-civ-spouse     Exec-managerial         Husband   White     Male   
2             Divorced   Handlers-cleaners   Not-in-family   White     Male   
3   Married-civ-spouse   Handlers-cleaners         Husband   Black     Male   
4   Married-civ-spouse      Prof-specialty            Wife   Black   Female   

   capital-gain  capital-loss  hours-per-week  native-country  income  
0          2174             0              40   United-States   <=50K  
```

**Task 1 Test**

<!-- grade_id: cell-788c5161452ee70a, points: 1, grade: true, solution: false -->
```python
# Task 1 - Test

# These tests are for you to check your work before submitting
assert isinstance(column_headers, list), 'Have you created a list of header names?'
assert len(column_headers) == 15, 'Did you include the correct number of items in the column_headers list?'
assert isinstance(adult, pd.DataFrame), 'Have you created a DataFrame named adult?'

# These tests will be completed when you submit your notebook for autograding
```

**Task 2** - Look at the first and last rows of the DataFrame

* Assign the first **ten (10)** rows of the `adult` DataFrame to `adult_head`
* Assign the last **ten (10)** rows of the `adult` DataFrame to `adult_tail`

<!-- grade_id: cell-7db4cedd667e6780, points: 0, grade: false, solution: true -->
```python
# Task 2

# YOUR CODE HERE
adult_head = adult.head(10)
adult_tail = adult.tail(10)

# Optional: print out adult_head and adult_tail
```

**Task 2 Test**

<!-- grade_id: cell-37f4eea6d7a1793b, points: 1, grade: true, solution: false -->
```python
# Task 2 - Test

# These tests are for you to check your work before submitting
assert isinstance(adult_head, pd.DataFrame), 'Have you created a DataFrame named adult_head?'
assert isinstance(adult_tail, pd.DataFrame), 'Have you created a DataFrame named adult_tail?'

# These tests will be completed when you submit your notebook for autograding
```

**Task 3** - Variable data types

For your `adult` DataFrame, determine the data types for the variables.

* Count the number of `int64` variable types and assign to `number_int64` (your value should be an integer)
* Count the number of `object` variable types and assign to `number_object` (your value should be an integer)

There are different ways to find the data types of the columns in a DataFrame. All you need to do for this task is to count the number of each data type and assign to the corresponding variable.

<!-- grade_id: cell-8f75ae5ed5309145, points: 0, grade: false, solution: true -->
```python
# Task 3

#print(adult.dtypes.value_counts().index)
number_int64 = adult.dtypes.value_counts()[1]
#number_object = adult.dtypes.value_counts()['object'] <- error
number_object = adult.dtypes.value_counts()[0]
#print(number_int64)
#print(number_object)
```

Output:
```
6
9
```

**Task 3 Test**

<!-- grade_id: cell-1425fb0c04a97cd3, points: 1, grade: true, solution: false -->
```python
# Task 3 - Test

# Hidden tests - you will see the results when you submit to Canvas
```

**Task 4** - DataFrame dimensions

* Find the dimensions of your DataFrame and assign result to the variable `adult_dimension`. Your variable should be a *tuple* and the row dimension should be listed first.

Hint: A tuple looks like this: (1, 2) - a tuple is a collection which is ordered and unchangeable. You can use the function `type(your_variable)` to print the variable type.

Hint 2: The `shape` method returns a tuple - convenient!

<!-- grade_id: cell-0dbc0cadefa896cd, points: 0, grade: false, solution: true -->
```python
# Task 4

# YOUR CODE HERE
adult_dimension = adult.shape

# Print out the shape
print(adult_dimension)
```

Output:
```
(32561, 15)
```

**Task 4 Test**

<!-- grade_id: cell-137d53dc191babc4, points: 1, grade: true, solution: false -->
```python
# Task 4 - Test

assert isinstance(adult_dimension, tuple), 'Have you created a tuple named adult_dimension?'

# Hidden tests - you will see the results when you submit to Canvas
```

**Task 5** - Missing values

Are there any missing values in the dataset? Let's check!

* Check for missing values using `.isnull().sum()`
* Count the number of missing values and assign the value to the variable `adult_missing`. Your variable should be an integer.

<!-- grade_id: cell-866428f0345c8e00, points: 0, grade: false, solution: true -->
```python
# Task 5

# YOUR CODE HERE
#print(adult.isnull().sum())
adult_missing = adult.isnull().sum().sum()
```

**Task 5 Test**

<!-- grade_id: cell-005d46af85b7d91b, points: 1, grade: true, solution: false -->
```python
# Task 5 - Test

# Hidden tests - you will see the results when you submit to Canvas
```

**Task 6** - Viewing DataFrame statistics

Let's look at some of the *summary statistics* of our dataset. We can use the .describe() function in order to see the summary statistics of the numeric columns.

Look at the statistics for the `adult` DataFrame and then complete the following two tasks:

* Find the value for the mean `age` and assign it to the variable `mean_age` (your value should be a float and defined to two decimal places)
* Find the standard deviation (std) for the `hours-per-week` variable and assign it to `std_hpw` (your value should be a float and defined to two decimal places)

Are there any values shown that might be a code for missing data?  We'll learn how to change the values in a DataFrame in the next module.

<!-- grade_id: cell-5dbc7144c66ecdec, points: 0, grade: false, solution: true -->
```python
# Task 6

# YOUR CODE HERE
#adult.describe()
mean_age = adult.describe()['age']['mean'].round(2)
std_hpw = adult.describe()['hours-per-week']['std'].round(2)
#print([mean_age, std_hpw])
```

**Task 6 Test**

<!-- grade_id: cell-8cd8fc43cc8628af, points: 1, grade: true, solution: false -->
```python
# Task 6 - Test

# Hidden tests - you will see the results when you submit to Canvas
```

**Task 7** - Summary statistics for non-numeric columns

We have some columns in this dataset that are non-numeric or object columns. This is usually a string object. Let's use the `describe()` function again but include the argument `exclude='number'` to exclude the numeric columns.

Using the results of `describe(exclude='number')`, complete the following two tasks:

* Find the number of unique `education` values and assign it to `unique_edu` (your value should be an integer)
* Find the number of times the most frequent observation for `income` occurs and assign it to `freq_income` (your value should be an integer) *(Note: this is not the income value itself, just how many times it occurs)*

<!-- grade_id: cell-7ed26707fcb247b0, points: 0, grade: false, solution: true -->
```python
# Task 7

# YOUR CODE HERE
#print(adult.describe(exclude='number'))
unique_edu = round(adult.describe(exclude='number')['education']['unique'])
freq_income = round(adult.describe(exclude='number')['income']['freq'])
#print(freq_income)
```

**Task 7 Test**

<!-- grade_id: cell-ce1a0a4faf19dbfc, points: 1, grade: true, solution: false -->
```python
# Task 7 - Test

# Hidden tests - you will see the results when you submit to Canvas
```

**Task 8** - Finding value counts

Let's look more specifically at the `relationship` column and perform a count of the number of observations for each category. We can see how many categories we have by using the `.unique()` method on the column. Then, we can use `.value_counts()` to count the number of observations in each category.

* View the unique values in the `relationship` column with `.unique()`
* View the number of observations for each value with `.value_counts()`
* Find the counts for `Other-relative` and assign to the variable `adult_other_rel`

<!-- grade_id: cell-c273f89ae164b596, points: 0, grade: false, solution: true -->
```python
# Task 8

# YOUR CODE HERE
#adult.describe(exclude='number')
adult['relationship'].unique()
adult['relationship'].value_counts()
adult['relationship'].value_counts().index
adult_other_rel = adult['relationship'].value_counts()[' Other-relative'] #leading whitespace, the invsible killer
#print(adult_other_rel)
```

Output:
```
981
```

**Task 8 Test**

<!-- grade_id: cell-7e8ae1c1c164dc92, points: 1, grade: true, solution: false -->
```python
# Task 8 - Test

# Hidden tests - you will see the results when you submit to Canvas
```

**Task 9** - Create a Series

A column of a DataFrame is a pandas Series. Using the `adult` DataFrame, create a Series from the `occupation` column.

* Create a Series from the `occupation` column and name it `adult_occup`

<!-- grade_id: cell-adbc8e3353d9e266, points: 0, grade: false, solution: true -->
```python
# Task 9

# YOUR CODE HERE
adult_occup = adult.get('occupation')
```

**Task 9 Test**

<!-- grade_id: cell-716ff49cd4887f23, points: 1, grade: true, solution: false -->
```python
# Task 9 - Test

# These tests are for you to check your work before submitting
assert isinstance(adult_occup, pd.Series), 'Have you created the correct Series?'

# Hidden tests - you will see the results when you submit to Canvas
```

## Feature Creation

Next you will focus on learning how to create new features from exisiting data. There are many ways to create features in your datasets; we'll only focus on a few here. Let's get started!

**Task 10** - Load the data set

* Using the provided URL, load the Ames Housing Data into a DataFrame called `house`
* View your DataFrame

<!-- grade_id: cell-76a11348650ab4bd, points: 0, grade: false, solution: true -->
```python
from pandas.io.parsers.readers import read_csv
# Task 10

# Imports
import pandas as pd

# Dataset URL

data_url = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_1/Housing/Ames%20Iowa%20Housing%20Data.csv'

# YOUR CODE HERE
house = read_csv(data_url)

# View your DataFrame
house.describe()
```

Output:
```
Id   MSSubClass  LotFrontage        LotArea  OverallQual  \
count  1460.000000  1460.000000  1201.000000    1460.000000  1460.000000   
mean    730.500000    56.897260    70.049958   10516.828082     6.099315   
std     421.610009    42.300571    24.284752    9981.264932     1.382997   
min       1.000000    20.000000    21.000000    1300.000000     1.000000   
25%     365.750000    20.000000    59.000000    7553.500000     5.000000   
50%     730.500000    50.000000    69.000000    9478.500000     6.000000   
75%    1095.250000    70.000000    80.000000   11601.500000     7.000000   
max    1460.000000   190.000000   313.000000  215245.000000    10.000000   

       OverallCond    YearBuilt  YearRemodAdd   MasVnrArea   BsmtFinSF1  ...  \
count  1460.000000  1460.000000   1460.000000  1452.000000  1460.000000  ...   
mean      5.575342  1971.267808   1984.865753   103.685262   443.639726  ...   
std       1.112799    30.202904     20.645407   181.066207   456.098091  ...   
min       1
```

**Task 10 Test**

<!-- grade_id: cell-851854e034a68c08, points: 1, grade: true, solution: false -->
```python
# Task 10 - Test

# These tests are for you to check your work before submitting
assert isinstance(house, pd.DataFrame), 'Have you created a DataFrame named house?'

# Hidden tests - you'll see results when you submit to Canvas
```

**Task 11** - Create a new feature

Now that we have the data loaded, we're going to create a new feature. We can use the some of the existing features, in this case, the total number of bathrooms in one column.

* Create a new variable called `Total_Bathrooms` that contains the total number of full and half bathrooms in the house.  

*Hint: Identify all the columns with bath in the title and add their values together. There are four columns to add together.  Add one to the total for each full or half bathroom.*

<!-- grade_id: cell-3ea6009bc1f45cc9, points: 0, grade: false, solution: true -->
```python
# Task 11

# YOUR CODE HERE
#print(house.columns)
house['Total_Bathrooms'] = house['FullBath'] + house['HalfBath'] + house['BsmtHalfBath'] + house['BsmtFullBath']
```

**Task 11 Test**

<!-- grade_id: cell-7505e6a0d5830f94, points: 1, grade: true, solution: false -->
```python
# Task 11 - Test

assert 'Total_Bathrooms' in house.columns, 'Did you add the new column?'

# Hidden tests - you'll see results when you submit to Canvas

# View all the bathroom columns
house[['Total_Bathrooms','BsmtFullBath','BsmtHalfBath','FullBath','HalfBath']].head()
```

Output:
```
Total_Bathrooms  BsmtFullBath  BsmtHalfBath  FullBath  HalfBath
0                4             1             0         2         1
1                3             0             1         2         0
2                4             1             0         2         1
3                2             1             0         1         0
4                4             1             0         2         1
```

**Task 12** - Create a new feature column

Next we're going to create another new column by "encoding" an existing column. Using the `LotArea` column, we'll assign values according the the area of the lot.

* Create a new column called `Large_Lot` that takes on the following values:
    * 0 when `LotArea` is **less** than 10,000
    * 1 when `LotArea` is **greater than or equal** to 10,000
    

<!-- grade_id: cell-f1cf383f1e791f0a, points: 0, grade: false, solution: true -->
```python
# Task 12

# YOUR CODE HERE
house.loc[house['LotArea'] >= 10000, 'Large_Lot'] = 1
house.loc[house['LotArea'] < 10000, 'Large_Lot'] = 0

# View the new column
#house[['LotArea', 'Large_Lot']].head()

```

**Task 12 Test**

<!-- grade_id: cell-ef0a2eeb1931e1dd, points: 1, grade: true, solution: false -->
```python
# Task 12 - Test

assert 'Large_Lot' in house.columns, 'Did you add the Large_Lot column?'

# Hidden tests - you'll see results when you submit to Canvas
```

**Task 13** - Create another feature

Let's continue with creating new features. This time, we'll focus on encoding for the sale month.

* Create a new column called `Summer_Sale` that takes on the following values:
    * 0 if the sale month was in September - May (`MoSold` = 1-5, 9-12)
    * 1 if the sale month was in June, July or August (`MoSold` = 6, 7 or 8)

<!-- grade_id: cell-8a1a021ee537cab5, points: 0, grade: false, solution: true -->
```python
# Task 13

# YOUR CODE HERE
summer_mask = (house['MoSold'] >= 6) & (house['MoSold'] <= 8) #bitwise and
house.loc[summer_mask, 'Summer_Sale'] = 1
house.loc[~ summer_mask, 'Summer_Sale'] = 0 #bitwise not
house['Summer_Sale'] = house['Summer_Sale'].astype('int') #the above line leaves them as floats due to NaN from first assignment


# View the new column
house[['MoSold', 'Summer_Sale']]
```

Output:
```
MoSold  Summer_Sale
0          2            0
1          5            0
2          9            0
3          2            0
4         12            0
...      ...          ...
1455       8            1
1456       2            0
1457       5            0
1458       4            0
1459       6            1

[1460 rows x 2 columns]
```

**Task 13 Test**

<!-- grade_id: cell-061d02cacd118c79, points: 1, grade: true, solution: false -->
```python
# Task 13 - Test

assert 'Summer_Sale' in house.columns, 'Did you add the Summer_Sale column?'

# Hidden tests - you'll see results when you submit to Canvas
```

```python

```